# CPSC 4300/6300-001 Applied Data Science (Fall 2020)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Wenkang Wei"
COLLABORATORS = "Wenkang Wei"
```

# CPSC4300/6300-001 Problem Set #1

In this problem set, you will create a NCAA College Football team dataset by scraping data from the Internet and perform some basic data manipulations. The purpose of this problem set is to get familiar with data collection with Python and gain some hand-on skills of using Python to some routine tasks.

The assignment has four parts: part A and B are auto-graded. Part C and D are manually graded by the TA. 

Before submission, make sure you run all the cells. Because the TA could not run your notebook, your grade will be only based on the output and the writing part of your submitted notebook. 

## Part A: Find the List of NCAA College Football Teams (30 points)

To start, we first create a list of NCAA football teams from ESPN's website. To save your time, we provide a Python module which scrapes the list of teams from url.

```python
%%writefile espn.py
import requests
from bs4 import BeautifulSoup
from requests.exceptions import HTTPError

NCAAF_URL="https://www.espn.com/college-football/teams"

def fetch_ncaaf_teams(url=NCAAF_URL):
    """Fetch the list of ncaa football teams
    
    Args:
        url (string): the url of the web page for ncaa football teams 
    
    Returns:
            
    """
    teams = []
    
    response = requests.get(url)
    html = response.content
    soup = BeautifulSoup(html, 'html.parser')
    
    conferences = soup.find_all("div", class_="headline")
    for conference in conferences:
        conference_name = conference.get_text()
        team_links = conference.parent.find_all("a", class_="AnchorLink")
        for team_link in team_links:
            if team_link.find("h2"):
                team_url = team_link['href']
                team_name = team_link.get_text()
                team_id = int(team_url.split("/")[-2])
                teams.append({"id": team_id, "name": team_name, "conference": conference_name, "url": team_url})
    return teams
```

Output:
```
Writing espn.py
```

```python
!ls
```

Output:
```
espn.py  part_a.ipynb  part_b.ipynb  part_c.ipynb  part_d.ipynb  __pycache__
```

### Question 1.  Import and use python functions

Write a Python statement to import the `fetch_ncaaf_teams()` function and use it to fetch the list of college football teams and save the list into a variable named `team_list`. (2 points)

<!-- grade_id: cell-51ce4adefe13d303, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from espn import fetch_ncaaf_teams
team_list = fetch_ncaaf_teams()
# show the first element in team_list
team_list[0]
#raise NotImplementedError()
```

Output:
```
{'id': 2132,
 'name': 'Cincinnati Bearcats',
 'conference': 'American',
 'url': '/college-football/team/_/id/2132/cincinnati-bearcats'}
```

```python
# test
assert type(team_list) is list
print("There are {} ncaa football teams".format(len(team_list)))
```

Output:
```
There are 129 ncaa football teams
```

<!-- grade_id: cell-7c8d7dff501a8d70, points: 1, grade: true, solution: false -->
```python
assert 'fetch_ncaaf_teams' in dir()
assert type(team_list) is list
```

<!-- grade_id: cell-efccda900d66ae1a, points: 1, grade: true, solution: false -->
```python
assert "Clemson Tigers" in [team["name"] for team in team_list]
```

### Question 2. Write files in JSON format

JSON (JavaScript Object Notation) is a common data file format that uses human-readable text to store and transmit data objects consisting of attribute–value pairs and array data types. For more, see https://en.wikipedia.org/wiki/JSON. Complete a code segment that write the team list into a JSON file named `"ncaaf_teams.json"`. (2 points) 

<!-- grade_id: cell-9bbddca1d06be714, points: 0, grade: false, solution: true -->
```python
import json
# YOUR CODE HERE
def save_teams(team_list, file_name = "ncaaf_teams.json"):
    team_list_json = json.dumps(team_list)
    fp = open("ncaaf_teams.json","w+")
    fp.write(team_list_json)
    fp.close()

save_teams(team_list)
#raise NotImplementedError()
```

<!-- grade_id: cell-ba507ed73c6aa940, points: 2, grade: true, solution: false -->
```python
import os
assert os.path.exists("ncaaf_teams.json")
```

## Question 3. Read JSON file into a Data Frame

Write a code segment that reads the teams from a JSON file into a DataFrame object named `teams`. (2 poinst)

<!-- grade_id: cell-c382e29c9e8a9601, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
# YOUR CODE HERE

def read_teams(file_name = "ncaaf_teams.json"):
    fp = open("ncaaf_teams.json",'r')
    team_list = json.load(fp)
    fp.close()
    team_list_df = pd.DataFrame(team_list)
    return team_list_df
teams = read_teams()
teams
# raise NotImplementedError()
```

Output:
```
id                     name conference  \
0    2132      Cincinnati Bearcats   American   
1     151    East Carolina Pirates   American   
2     248          Houston Cougars   American   
3     235           Memphis Tigers   American   
4    2426          Navy Midshipmen   American   
..    ...                      ...        ...   
124   309  Louisiana Ragin' Cajuns   Sun Belt   
125     6    South Alabama Jaguars   Sun Belt   
126   326      Texas State Bobcats   Sun Belt   
127  2653             Troy Trojans   Sun Belt   
128  2433       UL Monroe Warhawks   Sun Belt   

                                                   url  
0    /college-football/team/_/id/2132/cincinnati-be...  
1    /college-football/team/_/id/151/east-carolina-...  
2      /college-football/team/_/id/248/houston-cougars  
3       /college-football/team/_/id/235/memphis-tigers  
4     /college-football/team/_/id/2426/navy-midshipmen  
..                                                 ...  
124  /college-footb
```

<!-- grade_id: cell-2d7bc3c3a53a7cc4, points: 1, grade: true, solution: false -->
```python
assert type(teams) is pd.core.frame.DataFrame
```

<!-- grade_id: cell-2e40372204f3b359, points: 1, grade: true, solution: false -->
```python
assert pd.unique(teams["conference"]).size == 11
```

### Question 4. Examine a DataFrame Object

Write a code segment that outputs the first 20 rows of the Dataframe `teams` to a variable `row1_20`. (2 points)

<!-- grade_id: cell-5c80e40d49f47ef5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
row1_20 = teams[:20].copy()
# raise NotImplementedError()
print(row1_20.shape)
```

Output:
```
(20, 4)
```

<!-- grade_id: cell-52b5502600ae9cf7, points: 2, grade: true, solution: false -->
```python
assert row1_20.shape == (20, 4)
```

### Question 5. Find the columns of the DataFrame Object

Write a code segment that saves the names of the columns of the `teams` object. (2 points)

<!-- grade_id: cell-d9e0ccafd9bbe394, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
columns = teams.columns
columns
# raise NotImplementedError()
```

Output:
```
Index(['id', 'name', 'conference', 'url'], dtype='object')
```

<!-- grade_id: cell-8864b9f76b1645a4, points: 2, grade: true, solution: false -->
```python
assert type(columns) == pd.core.indexes.base.Index and columns.shape == (4,)
```

### Question 6. Find the size of the DataFrame Object

Write a single statement that outputs the number of rows and columns of the `teams` DataFrame Object and save the number rows to variable `nrow` and the number of columns to variable `ncols`. (2 points)

<!-- grade_id: cell-a2009e11d794ea2e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
nrows = teams.shape[0]
ncols = teams.shape[1]
# raise NotImplementedError()
nrows, ncols
```

Output:
```
(129, 4)
```

<!-- grade_id: cell-ba044b2c24a4ed96, points: 2, grade: true, solution: false -->
```python

```

### Question 7. Drop a Column in DataFrame

Write a statement to drop the url column in the teams DataFrame object without creating a new copy of the object. (2 points)

<!-- grade_id: cell-f9e65fc74aac24e7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
teams.drop(columns = ['url'], inplace= True)
# raise NotImplementedError()
```

<!-- grade_id: cell-9b0eac14466a2b27, points: 2, grade: true, solution: false -->
```python
assert teams.columns.size == 3 and "url" not in teams.columns
```

### Question 8. Set the team id as the Index of the teams object

<!-- grade_id: cell-b3d7345772bcbc66, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
teams.set_index("id", inplace=True)
# raise NotImplementedError()
```

```python
teams.head()
```

Output:
```
name conference
id                                    
2132    Cincinnati Bearcats   American
151   East Carolina Pirates   American
248         Houston Cougars   American
235          Memphis Tigers   American
2426        Navy Midshipmen   American
```

<!-- grade_id: cell-175e8e5b8e4b1537, points: 2, grade: true, solution: false -->
```python
assert(teams.index.name == 'id')
```

### Question 9. Select Rows by Column Value or Index

Select the team whose team id is 259 and save it team_259.

<!-- grade_id: cell-c6d40e192835a0c3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
team_259 = teams[teams.index==259]
# raise NotImplementedError()
team_259
```

Output:
```
name conference
id                                  
259  Virginia Tech Hokies        ACC
```

<!-- grade_id: cell-3b1fa2d6c5b0fd6b, points: 2, grade: true, solution: false -->
```python
assert team_259.iloc[0]['name'] == 'Virginia Tech Hokies'

```

### Question 10. Select Rows Based on Specific Conditions

Find rows in the teams DataFrame object in which the team in the ACC conference and save the results into a variable `acc_teams`. (2 points)

<!-- grade_id: cell-65ac6be1c0dd5249, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
acc_teams = teams[teams['conference']=='ACC']
# raise NotImplementedError()
acc_teams
```

Output:
```
name conference
id                                          
103         Boston College Eagles        ACC
228                Clemson Tigers        ACC
150              Duke Blue Devils        ACC
52        Florida State Seminoles        ACC
59    Georgia Tech Yellow Jackets        ACC
97           Louisville Cardinals        ACC
2390             Miami Hurricanes        ACC
152             NC State Wolfpack        ACC
153      North Carolina Tar Heels        ACC
87      Notre Dame Fighting Irish        ACC
221           Pittsburgh Panthers        ACC
183               Syracuse Orange        ACC
258            Virginia Cavaliers        ACC
259          Virginia Tech Hokies        ACC
154     Wake Forest Demon Deacons        ACC
```

```python
# test
assert pd.unique(acc_teams['conference'])[0] == 'ACC'
```

<!-- grade_id: cell-e3c5b1cd7cad993c, points: 1, grade: true, solution: false -->
```python
assert pd.unique(acc_teams['conference'])[0] == 'ACC'
```

<!-- grade_id: cell-f594e00c53b3bfeb, points: 1, grade: true, solution: false -->
```python
assert acc_teams['conference'].shape[0] == 15
```

### Question 11. Write a DataFrame Object to a JSON File

Write to the acc_teams object to a JSON file named `acc_teams.json`. (2 points)

<!-- grade_id: cell-8f47e341f2cdfe80, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def save_acc_teams(acc_teams, file_name = 'acc_teams.json'):
    fp = open(file_name, "w+")
    acc_teams_json = json.dumps(acc_teams.to_dict())
    fp.write(acc_teams_json)
    fp.close()
    
save_acc_teams(acc_teams)    
# pd.DataFrame(acc_teams.to_dict())
# raise NotImplementedError()

```

```python
!cat 'acc_teams.json'
```

Output:
```
{"name": {"103": "Boston College Eagles", "228": "Clemson Tigers", "150": "Duke Blue Devils", "52": "Florida State Seminoles", "59": "Georgia Tech Yellow Jackets", "97": "Louisville Cardinals", "2390": "Miami Hurricanes", "152": "NC State Wolfpack", "153": "North Carolina Tar Heels", "87": "Notre Dame Fighting Irish", "221": "Pittsburgh Panthers", "183": "Syracuse Orange", "258": "Virginia Cavaliers", "259": "Virginia Tech Hokies", "154": "Wake Forest Demon Deacons"}, "conference": {"103": "ACC", "228": "ACC", "150": "ACC", "52": "ACC", "59": "ACC", "97": "ACC", "2390": "ACC", "152": "ACC", "153": "ACC", "87": "ACC", "221": "ACC", "183": "ACC", "258": "ACC", "259": "ACC", "154": "ACC"}}
```

<!-- grade_id: cell-133c4949e798c9c8, points: 2, grade: true, solution: false -->
```python
assert os.path.exists('acc_teams.json')
```

### Question 12. The Data Oragnization of JSON Format

Compare the file `ncaaf_teams.json` and the file `acc_teams.json`, you may notice that even though both files store similar type of information, the data organizations in the two files are different. 

Which file organizes the data as a list of dict?

<!-- grade_id: cell-e076986e1ce3e639, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
filename = "ncaaf_teams.json"
# raise NotImplementedError()
filename
```

Output:
```
'ncaaf_teams.json'
```

<!-- grade_id: cell-83130931cdee051b, points: 2, grade: true, solution: false -->
```python

```

Suppose both files are read into memory and the number of teams is large, which data organization is more efficient to search a team by their team id?

<!-- grade_id: cell-b27d9f8ca645d79d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
filename = "acc_teams.json"
# raise NotImplementedError()
filename
```

Output:
```
'acc_teams.json'
```

<!-- grade_id: cell-bdb6f459cb3e79a9, points: 2, grade: true, solution: false -->
```python

```

### Question 13. CSV file format.

Write the `acc_teams` data frame into a csv file `acc_teams.csv`. (2 points)

<!-- grade_id: cell-7ec8fd00fd8d47fe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def save_cvs(acc_team, file_name="acc_teams.csv"):
    import csv
    csv_file = acc_team.to_csv()
    fp = open('acc_teams.csv','w+')
    fp.write(csv_file)
    fp.close()
save_cvs(acc_teams)

# raise NotImplementedError()
```

<!-- grade_id: cell-73b2723445b5d2a7, points: 2, grade: true, solution: false -->
```python
import os
assert os.path.exists('acc_teams.csv')
```

```python
!cat 'acc_teams.csv'
```

Output:
```
id,name,conference
103,Boston College Eagles,ACC
228,Clemson Tigers,ACC
150,Duke Blue Devils,ACC
52,Florida State Seminoles,ACC
59,Georgia Tech Yellow Jackets,ACC
97,Louisville Cardinals,ACC
2390,Miami Hurricanes,ACC
152,NC State Wolfpack,ACC
153,North Carolina Tar Heels,ACC
87,Notre Dame Fighting Irish,ACC
221,Pittsburgh Panthers,ACC
183,Syracuse Orange,ACC
258,Virginia Cavaliers,ACC
259,Virginia Tech Hokies,ACC
154,Wake Forest Demon Deacons,ACC
```

### Question 14. Sort Arrays

Sort the ACC teams by team names in alphabeta order and save the results in a list `sorted acc_team_names`. (2 points)

<!-- grade_id: cell-a9876762fdf2564b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import numpy as np
sorted_acc_team_names = np.sort(acc_teams['name'])
sorted_acc_team_names
# raise NotImplementedError()
```

Output:
```
array(['Boston College Eagles', 'Clemson Tigers', 'Duke Blue Devils',
       'Florida State Seminoles', 'Georgia Tech Yellow Jackets',
       'Louisville Cardinals', 'Miami Hurricanes', 'NC State Wolfpack',
       'North Carolina Tar Heels', 'Notre Dame Fighting Irish',
       'Pittsburgh Panthers', 'Syracuse Orange', 'Virginia Cavaliers',
       'Virginia Tech Hokies', 'Wake Forest Demon Deacons'], dtype=object)
```

<!-- grade_id: cell-da3dc40952d44f64, points: 2, grade: true, solution: false -->
```python
### BEGIN TESTS
assert ''.join([team[0] for team in sorted_acc_team_names]) == 'BCDFGLMNNNPSVVW'
### END TESTS
```

### Congratulations! You have finished Part A of the problem.