# CPSC 4300/6300-001 Applied Data Science (Fall 2020)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

# CPSC6300-002 Quiz 02

This quiz is based on problem 8 on Page 54 of the textbook "An Introduction to Statistical Learning". You may consult that page for a description of the `College.csv` dataset.

## Question 1 (1 points)

Write some python code to read the dataset "College.csv" into a pandas DataFrame object named `college`.

<!-- grade_id: cell-6e120e8b215b8c8b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cc2113d3435c6bf2, points: 1, grade: true, solution: false -->
```python
assert college is not None and college.shape[0] == 777
```

## Question 2 (2 Points)

Write some python code to change the name of the first column of the `college` DataFrame object to `CollegeName`. Your change must be made in place, i.e., the `college`s first column should be named as `CollegeName` after running your code.

<!-- grade_id: cell-5e5bf240ca3488a3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-53e2a685baf8b59f, points: 2, grade: true, solution: false -->
```python
assert college.columns[0] == "CollegeName"
```

## Question 3 (2 Points)

Write some python code to set the column "CollegeName" as the index of the `college` DataFrame object. The change must be made in place, i.e., a subsequent call `college.index.name` should output `CollegeName`.

<!-- grade_id: cell-f70e4d1cb9f7d2ce, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ee1d92397a6a22af, points: 1, grade: true, solution: false -->
```python
assert college.index.name == 'CollegeName'
```

<!-- grade_id: cell-188f626559256620, points: 1, grade: true, solution: false -->
```python
assert 'CollegeName' not in college.columns
```

## Question 4 (1 Points)

Write some python code to add a new column `AcceptRatio` into the college DataFrame object. The `AcceptRatio` is defined as $\frac{Number~of~applications~accepted}{Number~of~applications~received}$.

<!-- grade_id: cell-c71aa59197596c78, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1ed0a718d27afeee, points: 1, grade: true, solution: false -->
```python
assert "AcceptRatio" in college.columns and int(college.loc["Clemson University"]["AcceptRatio"] * 100) == 65
```

## Question 5 (2 Points)

Write some python code to compute the standard deviation of the graduation rate of non-private universities and save the value to a variable `grad_rate_std`.

<!-- grade_id: cell-0335856100be68a8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
grad_rate_std
```

<!-- grade_id: cell-4cacdb30cc3b3ba5, points: 2, grade: true, solution: false -->
```python

```

## Question 6 (2 Points)

Write some python code to split the university into two groups based on whether or not the proportion of students coming from the top 10% of their high school classes exceeds 50%. Save the numbers of the group of university with more than 50% students coming from 10% of their high school classes to a variable `n_elite_colleges`.

<!-- grade_id: cell-9fe2f7ba916c6410, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
n_elite_colleges
```

<!-- grade_id: cell-facb05a255f23a6f, points: 2, grade: true, solution: false -->
```python

```