# CPSC 4300/6300-001 Applied Data Science (Fall 2020)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

# CPSC4300/6300-001 Problem Set #3¶



# Part D. Develop Machine Learning Models (85 points)

```python
%matplotlib inline

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import matplotlib as mpt
import seaborn as sns

from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"
%matplotlib inline
```

## 1. Load the train and test dataset

In Part C, you have prepared a cleaned train data and test data. In this part, you develop machine learning models using the prepared data sets.

```python
df_train = pd.read_csv("input/housing_train_scaled.csv")
df_test = pd.read_csv("input/housing_test_scaled.csv")
```

## 2. Create a Data Structure to Keep Model Performance

In the model development phase, you may test numerous models. To compare these models, it is beneficial to create some utilities to keep the model performance data. Below is one example. 

```python
from collections  import defaultdict
from sklearn.metrics import mean_squared_error
from sklearn.metrics import explained_variance_score
from sklearn.metrics import mean_absolute_error
from sklearn.metrics import r2_score

model_perf = defaultdict(dict)

def compute_model_perf(model_name, tag, y_true, y_pred):
    model_perf[model_name]['{}.explained_variance'.format(tag)] = explained_variance_score(y_true, y_pred)
    model_perf[model_name]['{}.RSME'.format(tag)] = mean_squared_error(y_true, y_pred, squared=False)
    model_perf[model_name]['{}.R2'.format(tag)] = r2_score(y_true, y_pred)
    model_perf[model_name]['{}.RelativeError'.format(tag)] = mean_absolute_error(y_true/y_true, y_pred/y_true)
```

## 3. Build a base model

```python
from sklearn.linear_model import LinearRegression
X_train, y_train = df_train[df_train.columns[:-1]], df_train[df_train.columns[-1]]
X_test, y_test = df_test[df_test.columns[:-1]], df_test[df_test.columns[-1]]
```

__Question 3(a).__ Write code to train a linear regression model that uses all the features, save your model to a variable named `base_model`.

<!-- grade_id: cell-392cbf512e706546, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
base_model.coef_, base_model.intercept_, len(X_train.columns)
```

<!-- grade_id: cell-bd9c9f368d79e6cd, points: 1, grade: true, solution: false -->
```python
isinstance(base_model, LinearRegression)
```

<!-- grade_id: cell-0620c0eedf9c5dff, points: 2, grade: true, solution: false -->
```python
assert isinstance(base_model.coef_, np.ndarray)
assert base_model.coef_.shape[0] == X_train.shape[1]
```

__Question 3(b).__ Write code to compute the compute the following metrics for both the train and test data.

+ explained_variance
+ root_mean_squared_error
+ r2
+ mean_relative_error

Add the results to the dictionary `model_perf` with `LinearRegressionBase` as the model_name and `Train` and `Test` for the tags of train fit performance and test fit performance respectively.

Hint: you can use the `compute_model_perf` utility.

<!-- grade_id: cell-cda7ae1cbff6b325, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
df_perf = pd.DataFrame(model_perf).T
df_perf
```

<!-- grade_id: cell-1afbb00b13da4079, points: 4, grade: true, solution: false -->
```python
tags = np.unique([col.split('.')[0] for col in df_perf.columns])
assert 'Train' in tags and 'Test' in tags
```

### Analyze modeling errors

Besides the performance metrics we collected above, we can also plot the errors to detect some patterns in the model prediction errors. 

```python
fig, ax = plt.subplots(1, 2, figsize=(14, 6))
_ = sns.scatterplot(x=y_train, y=(y_train-y_train_pred), ax = ax[0], label='Train')
_ = sns.scatterplot(x=y_test, y=(y_test-y_test_pred), ax = ax[0], label='Test')
ax[0].set_ylabel('Absolute Error')
_ = sns.scatterplot(x=y_train, y=(1.0-y_train_pred/y_train), ax = ax[1], label='Train')
_ = sns.scatterplot(x=y_test, y=(1.0-y_test_pred/y_test), ax = ax[1], label='Test')
_ = ax[0].set_ylabel('Relative Error')
```

__Question 3(c).__ Based on the above and model performance results and plots, do you think linear regression is a biased model or not? Justify your conclusion. (6 points: 2 point for answer; 4 points for justification)

### Type your answer here
1. Is the base linear regression model  biased?
2. Justify your conclusion

__Question 3(d)__. List one pattern that you have observed in the model prediction errors and describe in which way you may be able to reduce the prediction errors of the base linear regression model. (6 points: 2 points for pattern; 4 points for ways of improvement) 

### Type your answer here
1. Patterns observed
2. Ways of improvements

## 4. Hypothesis Tests on the Coefficients

__Question 4(a).__ Estimate an ordinary least squares (OLS) linear regression model using the `statsmodels` package and output the model parameters. (4 points)

<!-- grade_id: cell-dcefe9ded470a484, points: 4, grade: true, solution: true -->
```python
import statsmodels.api as sm
import statsmodels.formula.api as smf

# YOUR CODE HERE
raise NotImplementedError()
base_ols_model.summary().tables[1]
```

__Question 4(b).__ Based on the t-statistic and p-value in the output of 4(a), which feature would be most irrelevant in the base regression model? Explain why? (6 points: 2 points for identify the feature; 4 points for explanation)

### Type your answer here
1. Irrelevant feature identified
2. Explain

## 5. Feature Selection

Feature selection is an important component in model development. In this question, we apply univariate F-test statistics and mutual information to the housing data.

```python
from sklearn.feature_selection import f_regression
from sklearn.feature_selection import mutual_info_regression

f_test, pval = f_regression(X_train, y_train)

mi = mutual_info_regression(X_train, y_train)

df_features = pd.DataFrame({'f_test': pd.Series(f_test, index=X_train.columns), 
                   #'pval': pd.Series(pval, index=X_train.columns),
                   'mi': pd.Series(mi, index=X_train.columns)
                  })
df_features['f_test'] = df_features['f_test']/df_features['f_test'].max()
df_features['mi'] = df_features['mi']/df_features['mi'].max()
```

```python
df_features.sort_values(by='mi', ascending=False)
```

```python
df_features.sort_values(by='mi', ascending=True)[['f_test', 'mi']].plot.barh()
```

#### __Question 5(a).__ Brief describe what F-Test in regression is? (4 points)

### Type your answer here
what is F_Test?

__Question 5(b).__ Write some code to find the top 5 most discriminative features according to their F-Test scores? (3 points)

<!-- grade_id: cell-f0200ab1937a3b64, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

__Question 5(c)__ Modify the following code to assess the linear regression models that use the set of discriminative features sorted by the feature's __F-Test__ scores and __mutual information (MI)__ scores. (5 points)

```
y =  df_train[df_train.columns[-1]]
y_test = df_test[df_test.columns[-1]]
model_perf = defaultdict(dict)

metric = 'f_test'
for n_features in range(2, len(df_train.columns)):
    features = list(df_features.sort_values(by=metric, ascending=False).iloc[:n_features].index)
   
    model_name = metric + "-{}".format(n_features)
    model = LinearRegression()
    
    X = df_train[features]
    _ = model.fit(X, y)    
    y_train_pred = model.predict(X)    
    
    X_test = df_test[features]
    y_test_pred = model.predict(X_test)
    
    compute_model_perf(model_name, 'Train', y, y_train_pred)
    compute_model_perf(model_name, 'Test', y_test, y_test_pred)
```

<!-- grade_id: cell-756e1cd7e24e3813, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
df_perf = pd.DataFrame(model_perf).T

index = pd.MultiIndex.from_tuples([(a[0], int(a[1])) for a in df_perf.index.str.split('-')])
df2 = df_perf.set_index(index)
df_r2 = pd.DataFrame({'f_test': df2.loc['f_test', 'Test.R2'], 'mi': df2.loc['mi', 'Test.R2']})
fig, ax = plt.subplots(1, 1, figsize=(8, 6))
_ = df_r2.plot.bar(ax=ax)
_ = ax.set_xlabel('n_features')
_ = ax.set_ylabel('Test.R2')
_ = ax.set_ylim(0.4, 0.7)
_ = ax.set_title('Test.R2 Under Different Set of Features')
```

## 6. Nonlinear Regression Models

__Question 6(a).__ The error plots in question 3 indicate a non-linear relationship between the features and the target variable. Complete the following code to perform a polynomial regression. (5 points)

<!-- grade_id: cell-ff284c9999260a56, points: 5, grade: true, solution: true -->
```python
from sklearn.preprocessing import PolynomialFeatures
from sklearn.pipeline import make_pipeline

y =  df_train[df_train.columns[-1]]
y_test = df_test[df_test.columns[-1]]
model_perf = defaultdict(dict)

for deg in range(1, 4):
# YOUR CODE HERE
raise NotImplementedError()
df_perf = pd.DataFrame(model_perf).T
```

```python
df_perf = pd.DataFrame(model_perf).T
df_perf
```

__Question 6(b)__. In the above results, `Train.R2` descreases as the order of the polynomial increases. Does that mean it is beneficial to use a higher order polynomial in the regression? Justify your answer. (6 points: 2 points for correct answer; 4 points for justification)

### Type your answer here
1. Is a higher order polynomial preferable?
2. Justification

## 7. K-Nearest Neighbors Regression

The k-nearest neighbors (___k_-NN__) algorithm is a type of instance-based learning where the function is only approximated locally and all computation is deferred until function evaluation.

__Question 7(a)__. Complete the following code to assess the performance of ___k_-NN__ regression on the housing data. (5 points)

<!-- grade_id: cell-ed6617a45bf6264d, points: 5, grade: true, solution: true -->
```python
from sklearn.neighbors import KNeighborsRegressor

model_perf = defaultdict(dict)

y = df_train[df_train.columns[-1]]
y_test = df_test[df_test.columns[-1]]

features = df_train.columns[:-1]
X = df_train[features]
X_test = df_test[features]

for n_neighbors in range(1, 16):
# YOUR CODE HERE
raise NotImplementedError()
```

```python
df_perf = pd.DataFrame(model_perf).T
df_perf
```

```python
fig, ax = plt.subplots(1, 1, figsize=(8, 6))
_ = df_perf[['Train.R2', 'Test.R2']].plot.bar(ax=ax)
_ = ax.set_xlabel('n_neighbors')
_ = ax.set_ylabel('Test.R2')
_ = ax.set_title('Test.R2 Under Different Set of Features')
```

__Question 7(b)__. What pattern(s) have you observed from the above results? Explain why such pattern or patterns occur. (6 points: 2 points for pattern; 4 points for explanation)

### Type your answer here
1. Pattern
2. Explain

__Question 7(c)__. In the above study, we use the scaled data set. Assume we use a non-scaled data set in ___k_-NN__ regression. How do you expect the model performance change? Explain why.  (6 point: 2 points for possible change; 4 points for explanation)

### Type your answer here
1. Change
2. Explain

## 8. Random Forest Regression

Random forests regression is an ensemble learning method that operates by constructing multiple decision trees and outputs the  mean/average prediction of the individual trees.

__Question 8(a)__. Complete the following code to assess the performance of Random Forest Regression on the housing data. (5 poinst)

<!-- grade_id: cell-25e49cfb4f0eac2b, points: 5, grade: true, solution: true -->
```python
from sklearn.ensemble import RandomForestRegressor

model_perf = defaultdict(dict)

y = df_train[df_train.columns[-1]]
y_test = df_test[df_test.columns[-1]]

features = df_train.columns[:-1]
X = df_train[features]
X_test = df_test[features]

for n_estimators in [10, 50, 100, 200]:
# YOUR CODE HERE
raise NotImplementedError()
```

```python
df_perf = pd.DataFrame(model_perf).T
df_perf
```

__Question 8(b)__. Brief describe any observation(s) you have regarding the model performance of the Random Forest Regression. (6 Points)

### Type your answer here

### Model Persistence

Model development is a time consuming process. When you have developed a good model, you should save the model and load it for prediction.

```python
# Create a model 
from sklearn.ensemble import RandomForestRegressor

features = df_train.columns[:-1]
X = df_train[features]
y = df_train[df_train.columns[-1]]

model_rf200 = RandomForestRegressor(n_estimators=200)
model_rf200.fit(X, y)
```

```python
# Save the model using joblib.dump(); to load the saved model, use joblib.load()
import joblib
joblib.dump(model_rf200, "models/model_rs200.joblib")
```

```python
# Load the predict data
df_to_predict = pd.read_csv("input/housing_prediction_scaled.csv")
df_to_predict
```

__Question 8(c)__. Load the saved model `models/model_rs200.joblib` and apply the model to predict the median housing price of some unknown blocks stored in `../input/housing_prediction_scaled.csv`. Save your results to y_predict. (5 points)

<!-- grade_id: cell-1d11f074f89ff1cc, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

If you are interest in seeing the performance of your model, you can peek the actual value as follows.

```python
y_actual = pd.read_csv("input/housing_prediction_scaled_y.csv")['median_house_value']
df_predict = pd.DataFrame({'Actual': y_actual, 'Predict': y_predict})
df_predict['Absolute_Error'] = df_predict['Predict'] - df_predict['Actual']
df_predict['Relative_Error'] = df_predict['Absolute_Error'] / df_predict['Actual']

df_predict
```

```python
# Remove the model files since it is too big
import os
os.remove("models/model_rs200.joblib")
assert not os.path.exists("models/model_rs200.joblib")
```

__Congratulations! You have done such a wonderful job that your model actually beats the experts. That should make your boss happy and consider offering you a raise.__